# Project 1

## Assignment Instructions

For this project we'll have relatively detailed instructions. For the next project, we'll have more general instructions, since by that time you'll be more comfortable using the various programming tools.
Be sure to tackle this program in stages. Get one thing working correctly, and then move on to the next item. These instructions are written in the order that I recommend you follow.  

First, let’s get a menu going.  This menu will have 4 options.  The first 3 menu choices will run different tests, and the 4th option will quit the program.  (Notice that the 3rd option is for a function called randu().  We’ll learn later that this is an optional part of the assignment – but for consistency, you are required to have the menu option present in your code.)

## Sample Output

### Sample Output 1

When you run your program, it should produce output similar to what is shown below.  It does not need to be identical, but should be functionally the same.
<pre>
<code>
Sample output 1
Choose from the following options: 

1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 

Enter your choice:  1 

Testing the built-in python random number generator. 

Enter number of tests:  100 


Digit Number of times generated
  0     11
  1     5
  2     16
  3     8
  4     14
  5     8
  6     7
  7     11
  8     8
  9     12
 

Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 

Enter your choice:  4 

Goodbye!
</code>
</pre>

### Sample Output 2

When you run your program, it should produce output similar to what is shown below.  It does not need to be identical, but should be functionally the same.
<pre>
<code>
Sample output 2 

Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 
 

Enter your choice:  1 
 

Testing the built-in python random number generator. 
 

Enter number of tests:  1000000 
 

Digit Number of times generated
  0     99926
  1     99424
  2     100102
  3     100584
  4     99864
  5     99842
  6     99874
  7     99913
  8     100272
  9     100199


 

Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the Dilbert() function 
3. Test the Randu() function. 
4. Quit this program 
 

Enter your choice:  2 
 

Test Dilbert's random number generator. 
 
Enter number of tests:  1000000 
 
Digit Number of times generated
  0     0
  1     0
  2     0
  3     0
  4     0
  5     0
  6     0
  7     0
  8     0
  9     1000000

Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 
 

Enter your choice:  4 
 

Goodbye!
</code>
</pre>

### Sample Output 3

When you run your program, it should produce output similar to what is shown below.  It does not need to be identical, but should be functionally the same.
<pre>
<code>
Sample output 3 (Optional RANDU function) 

Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 
 

Enter your choice:  3 
 

Test the RANDU random number generator. 
 

Enter number of tests:  1000000 
 

Digit Number of times generated
  0     0
  1     200251
  2     0
  3     199253
  4     0
  5     200147
  6     0
  7     200339
  8     0
  9     200010

 
Choose from the following options: 
1. Test the python random.randrange() function 
2. Test the dilbert() function 
3. Test the randu() function. 
4. Quit this program 
 

Enter your choice:  4 
 

Goodbye!
</code>
</pre>

## Let's start coding!

```python
# Install necessary dependencies
import random
import matplotlib.pyplot as plot
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[2], line 3
      1 # Install necessary dependencies
      2 import random
----> 3 import matplotlib.pyplot as plot

ModuleNotFoundError: No module named 'matplotlib'
```

### Part A: Set up the menu

1. Reuse your previous code from Programming Assignment 4 that creates a menu, and modify it to display these menu options, and then return the number selected:


<ol style="padding-left: 100px;">
    <li>Test the python random.randrange()</li>    
    <li>Test the dilbert() function</li>
    <li>Test the randu() function</li>   
    <li>Quit this program  
</ol>

2. Add a while loop, so that if the user chooses an incorrect item, they see the error message 'That is not a valid option. Please try again.'  (Hint:  you can use a ‘while True:’ loop, and then have an if statement that contains a ‘break’ if they choose option 4.)

3. Once this is working--save a backup copy of your code somewhere on your computer! (Download the ipynb file, and save it on your device). It's very common to break your old code when you add new code, so you want to save a copy anytime you get something working right!

#### Write the menu() function here, and be sure to manually test it:

<!-- grade_id: cell-3d969c25ce3dd871, points: 0, grade: false, solution: true -->
```python
def menu():
    """Display a menu with four options, return the user’s selection as an integer"""
    # YOUR CODE HERE
    raise NotImplementedError()


# Call your function here to test it.  Does it work correctly?
print(f"The menu function returned {menu()}.")
```

#### Now, test your code with some automated tests:

Once you have written the menu() function above and completed your manual testing, you should try the auto-grader tests in the next cell.  The tests below simulate how a user would select each menu option.  For example, the first test pretends that the user enters the character “1”, and verifies that your menu() function returns the integer 1.

Note that these tests are very close to the actual tests used by the auto-grader, so if your menu() function is passing these tests (so that you see no errors output when you run the cell), your code is likely correct.

__NOTE: If this cell causes an infinite loop click the square icon in the tray above to quickly kill the kernel.  If an infinite loop happens that means that your menu code is not correct and the tests are failing.  When this cell runs without error, your code has pass the auto-tests.__

<!-- grade_id: cell-efada7f6d48c7b35, points: 10, grade: true, solution: false -->
```python
### BEGIN TESTS
from unittest.mock import patch

def test_menu():
    # Test for input '1'
    with patch('builtins.input', return_value="1"):
        assert menu() == 1, "menu() cannot handle the input of 1 properly, check logic"
        print("USER ENTERS: 1")
        print("RETURNED value of int(1) for answer\n\n")

    # Test for input '2'
    with patch('builtins.input', return_value="2"):
        assert menu() == 2, "menu() cannot handle the input of 2 properly, check logic"
        print("USER ENTERS: 2")
        print("RETURNED value of int(2) for answer\n\n")
        
    # Test for input '3'
    with patch('builtins.input', return_value="3"):
        assert menu() == 3, "menu() cannot handle the input of 3 properly, check logic"
        print("USER ENTERS: 3")
        print("RETURNED value of int(3) for answer\n\n")
        
    # Test for input '4'
    with patch('builtins.input', return_value="4"):
        assert menu() == 4, "menu() cannot handle the input of 4 properly, check logic"
        print("USER ENTERS: 4")
        print("RETURNED value of int(4) for answer\n\n")

    # Test for an invalid input followed by a valid input
    print("simulating a user entering an invalid value, followed by a valid value")
    with patch('builtins.input', side_effect=["5", "1"]):
        assert menu() == 1
        print("USER ENTERS: 1")
        print("RETURNED value of int(1) for answer\n\n")
        
    
test_menu()
    
### END TESTS
```

### Parts B-C: Test and graph random.randrange()

#### Part B: Test Python's random.randrange() function

1.  Eventually, we want to write a `main()` so that if the user chooses option 1, they will test the python’s built-in function `random.randrange()`.  If they choose that option, we want to call a function named `test_random_randrange()`.  This function will then proceed to generate many random numbers, and test to see if they are roughly evenly distributed.  In order to do this, let’s follow these steps below, and first write the function `test_random_randrange()`.  

2.  Write a function named test_random_randrange() which will:

    - take one argument -- the number of tests to conduct, and return nothing (it will just print the results) 

    - set up 10 counter variables that will be used to count how many of each digit is generated, for example:  `number_of_0s = 0`, `number_of_1s = 0`, and so on.  You probably want to choose a shorter variable name, though.

    - set up a for loop that will loop to repeat a bunch of tests (the number that was passed into the function).  In each loop: 

        - Generate a new random number, using the `random.randrange(0,10)` function, so that we get a number between 0 and 9 inclusive 

        - Use an  if statement to increment the correct counter (i.e. if you get a `1`, add one to the `number_of_1s` variable) 

    - after the for loop exits, print out the total counts for each digit (see sample output).

        - __IMPORTANT NOTE__: The autograder will check to see that you have a descriptive header row printed, that explains the two columns, so you will need to include a header that contains the words “Digit” and “Number”.  It will also check that all results for each digit are printed.
        - __HINT__: to make this look nice, you can use a format string, with the alignment options, for example >10, as discussed in lecture, to get things to line up nicely in a column.
        - __ALSO__: this is a pseudorandom number generator, so each time you run the code the numbers will be different (and that is ok).

<div style="border-bottom: 100px;"></div>

3.  Later on, towards the bottom of this Jupyter Lab, we’ll call this function from main().  But for now, let’s do some manual and automated testing on this function by itself.  Use the code cell that is shown under Part C (where we will add some graphing to this code) to write the function, as well as test what you’ve written.  

4. Run this several times, and fix any bugs that you find, until you're satisfied with it 

5. Save a backup copy on your computer! 

#### Part C: Make a graph to show the data for random.randrange()

1. There is a great package named matplotlib that will allow us to generate a bar chart in 8 lines of code.  Since we haven't encountered it yet in the lecture, I'll give you the code.

2.  Add these lines of code to the end of the function that you wrote in part B.  
>_the variables_ `n0`, `n1` _etc should correspond to what you named your_ `number_of_0s` _variables in step B._

      
       The code is:
```python
    # display the results as a graph 
    figure = plot.figure()  # start with a blank figure 
    axes = figure.add_axes([0,0,1,1]) 
    labels = ['0','1','2','3','4','5','6','7','8','9'] 
    values = [n0,n1,n2,n3,n4,n5,n6,n7,n8,n9] 
    axes.bar(labels,values) 
    axes.set_title('Counts: random.randrange()') 
    plot.show()
```
3. Compare your output to the sample output, and save a backup of your code

#### Write the test_random_randrange() function here, and be sure to manually test it:

<!-- grade_id: cell-7288adaf5d582fed, points: 0, grade: false, solution: true -->
```python
def test_random_randrange(num_to_test):
    # YOUR CODE HERE
    raise NotImplementedError()

    
# Call your function here to test it.  Does it work correctly?
test_random_randrange(25)
```

#### Once you have finished both Parts B and C and manually tested them, test your code with automated tests here.:

This code will test the `test_random_randrange` function with different values of `num_to_test`, ensuring that the function's behavior is consistent across different inputs. It also includes a basic test for the plot, checking that the total value of the bars in the plot matches `num_to_test`.

Note: 
If your function is not quite right, you’ll see some errors output. You can use the error messages for guidance as you adjust your function. If your code is passing, you will see no errors, and instead you will see the output of several plots corresponding to different numbers of random tests that were generated.

<!-- grade_id: cell-43e6488137371dbc, points: 25, grade: true, solution: false -->
```python
### BEGIN TESTS
import random
import matplotlib.pyplot as plot
from unittest.mock import patch

def autotest_random_randrange():
    # Test with different values of num_to_test
    for num_to_test in [0, 1, 10, 100, 1000]:
        from io import StringIO
        import sys

        captured_output = StringIO()
        sys.stdout = captured_output


        with patch('matplotlib.pyplot.show'): 
            test_random_randrange(num_to_test) 

        sys.stdout = sys.__stdout__

        output_lines = captured_output.getvalue().strip().split('\n')


        assert len(output_lines) == 11, f"Expected 11 lines in the output, but got {len(output_lines)} lines."

       
        header_lower = output_lines[0].lower().strip()
        target_words = ["digit", "number"]
        
        assert all(word in header_lower for word in target_words), "The header is missing one or more required words ('digit', 'number'). Case does not matter."


     
        total_count = 0
        for line in output_lines[1:]:
            digit, count = line.split()
            digit = int(digit)
            count = int(count)


            assert digit in range(10), f"Expected digit between 0 and 9, but got {digit}"

            assert 0 <= count <= num_to_test, f"Expected count between 0 and {num_to_test}, but got {count}"

            total_count += count

 
        assert total_count == num_to_test, f"Expected total count to be {num_to_test}, but got {total_count}"


        figure = plot.gcf()
        axes = plot.gca()
        bars = axes.patches
        legend = axes.legend([f"num_to_test: {num_to_test}"])
        bar_values = [bar.get_height() for bar in bars]
        assert sum(bar_values) == num_to_test, f"Expected total bar values to be {num_to_test}, but got {sum(bar_values)}"

autotest_random_randrange()
### END TESTS
```

### Part D: Repeat parts B and C for dilbert():

1.  For this part, we will write two functions.  First, let’s create our own random number generator that is actually pretty terrible.  Create a new function in the appropriate cell below, named dilbert() that always returns the number 9.

2.  Second, write a function that will test this (bad) random number generator.  Let’s name it test_dilbert().  It should do all the same things as the test_random_randrange() function did.  In fact, you’ll probably copy/paste a lot of the code and then modify it.  (Note that, usually, when we copy/paste code, it’s often a good idea to create a new function instead…but in this case we want to just focus on the goals of the project, so I do recommend copy/paste this time.)

3.  Test your code and compare your results to the sample output above.

4.  Save a backup of your code!

#### Write your dilbert() function here:

<!-- grade_id: cell-31508445ded5bb52, points: 0, grade: false, solution: true -->
```python
def dilbert():
    # YOUR CODE HERE
    raise NotImplementedError()
```

#### Write the test_dilbert() function here, and be sure to manually test it:

<!-- grade_id: cell-197a43ae9982ade0, points: 0, grade: false, solution: true -->
```python
def test_dilbert(num_to_test):
    # YOUR CODE HERE
    raise NotImplementedError()

    
# Call your function here to test it.  Does it work correctly?
test_dilbert(25)
```

#### Once you have finished Part D and manually tested your functions, test your code with automated tests here:

This code will test the `dilbert` function and then the `test_dilbert` function with different values of `num_to_test`, ensuring that the function's behavior is consistent across different inputs. It also includes a basic test for the plot, checking that the total value of the bars in the plot matches `num_to_test`.

___Note: When this cell fails, you must adjust your function. You can use the traceback here for guidance as you adjust and continue to code out your function.  If your code is passing, you will see no errors, and instead you will see the output of several plots corersponding to the randomly generated `num_to_test` value.__

<!-- grade_id: cell-b042ff8f61e5d0d0, points: 35, grade: true, solution: false -->
```python
### BEGIN TESTS
import matplotlib.pyplot as plot
from unittest.mock import patch

def autotest_dilbert():
    for num_to_test in [0, 1, 10, 100, 1000]:
        from io import StringIO
        import sys    
        captured_output = StringIO()
        sys.stdout = captured_output
        
        assert dilbert() == 9, f"the dilbert function returned: {dilbert()}, but was expected to return 9 (data type: int)"
        
        with patch("matplotlib.pyplot.show"):
            test_dilbert(num_to_test) 

        sys.stdout = sys.__stdout__  

        output_lines = captured_output.getvalue().strip().split('\n')

        assert len(output_lines) == 11, f"Expected 11 lines in the output, but got {len(output_lines)} lines."
        
        header_lower = output_lines[0].lower().strip()
        target_words = ["number", "digit"]
        assert all(word in header_lower for word in target_words), "The header is missing one or more required words ('digit', 'number'). Case does not matter."


        for line in output_lines[1:-1]:
            digit, count = line.split()
            digit = int(digit)
            count = int(count)


            assert digit in range(9), f"Expected digit between 0 and 8, but got {digit}"
            assert count == 0, f"Expected count to be 0, but got {count}"


        digit, count = output_lines[-1].split()
        assert int(digit) == 9, f"Expected digit to be 9, but got {digit}"
        assert int(count) == num_to_test, f"Expected count to be {num_to_test}, but got {count}"

    
        figure = plot.gcf()
        axes = plot.gca()
        bars = axes.patches
        legend = axes.legend([f"num_to_test: {num_to_test}"])
        bar_values = [bar.get_height() for bar in bars]
        assert bar_values == [0] * 9 + [num_to_test], f"Expected bar values to be {[0] * 9 + [num_to_test]}, but got {bar_values}"
        

autotest_dilbert()
### END TESTS
```

### Part E: Pulling it all together in main()

- Now that we’ve written and tested the functions, let’s pull it all together into a main() function that lets the user interact with the program.  We can break this down into steps as well.

- First, we want to let the user choose which function they want to test (using the menu() function).

- Once they make a choice, we want to ask them now many numbers they want to test.

- Now, we can call the selected test function (test_random_randrange(), test_dilbert(), or the optional test_randu()) and pass in the number of tests requested.

- Finally, we want to repeatedly loop until the user wants to quit.

#### Write your main() function so that it meets the requirements above:

<!-- grade_id: cell-1a1ecf64975af561, points: 0, grade: false, solution: true -->
```python
def main():
    # YOUR CODE HERE
    raise NotImplementedError()
```

#### Execute this cell to call your main function, and test it thoroughly.

```python
# Execute this cell to call your main function, and test it thoroughly.
main()
```

#### An important note on autotesting the program associated with the main() function:

The decisions we make while solving problems with programming are similar to many of the decisions we make in life, often with no one correct answer, but instead require analysis of the tradeoffs among the possible answers and careful consideration of the context or environment they exist in.  Testing code is no stranger to this, and with that said, autotesting the function associated with our main program comes down to the tradeoff between the complexity of trying to automate the multiple side effects associated with a number of functions that it coordinates the execution for (`random.randrange`, `dilbert`, `randu` ... ) versus the easy and non-complex manual testing that could be done given the relatively simple nature of the main program's code itself.  Therefore, the course facilitators will manually grade your main() function, and ensure that it meets the requirements.  The rubric (grading guide) used by the course facilitators is shown here.  Please check that you have met all the requirements below – then follow the instructions [here](https://www.coursera.org/learn/ball-state-univeristy-introduction-to-programming/assignment-submission/QgjpS/submit-project-1) to submit your final version of the project.

### Extension: RANDU

If you’re up for a challenge, there is another function that is very interesting to test.  It is known as RANDU, or "the most ill-conceived random number generator ever designed" according to Wikipedia. 

In order to write solve this part, you'll need to: 

1.  Write the RANDU function.  It takes one argument (the last generated value), and returns the next number according to the formula: `65539 * last_number % pow(2, 31)`

2.  In the function that tests RANDU, you will need to start that "last" number with a value of `1`. 

3.  Then, after RANDU returns the number, you'll need to FIRST store it into that "last" variable, THEN take the `% 10` of that number (to just look at the last digit), and use that when you are counting the distribution of random numbers. 

_Hint:  Yes, RANDU is really that bad.  It only generates odd numbers, never an even number._

#### Code RANDU:

<!-- grade_id: cell-f7fa1004495158b6, points: 0, grade: false, solution: true -->
```python
def randu(last_number):
    # YOUR CODE HERE
    raise NotImplementedError()
```

#### Test RANDU manually:

<!-- grade_id: cell-f16643b1ca14bd45, points: 0, grade: false, solution: true -->
```python
def test_randu(num_to_test):
    # YOUR CODE HERE
    raise NotImplementedError()
```